# Week 12-1 Practice

### Recap

Last class we implmented functions called `f_torch` that evaluated $f(x) = 3x^\intercal x - x_1 - 4$.

### Problem 1: Minimizing our function with gradient descent
Write a function that minimizes $f$ using gradient descent, by calling by calling `f_torch`. The function should take two arguments: an initial iterate `x0` and a number of gradient descent iterations `num_iters`. It should return the minimizer.

```python
!pip install "numpy<2.0"
```

Output:
```
Collecting numpy<2.0
  Downloading numpy-1.26.4-cp312-cp312-macosx_10_9_x86_64.whl.metadata (61 kB)
Downloading numpy-1.26.4-cp312-cp312-macosx_10_9_x86_64.whl (20.3 MB)
[2K   [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m20.3/20.3 MB[0m [31m24.3 MB/s[0m eta [36m0:00:00[0m00:01[0m00:01[0m
[?25hInstalling collected packages: numpy
  Attempting uninstall: numpy
    Found existing installation: numpy 2.0.1
    Uninstalling numpy-2.0.1:
      Successfully uninstalled numpy-2.0.1
[31mERROR: pip's dependency resolver does not currently take into account all the packages that are installed. This behaviour is the source of the following dependency conflicts.
transformers 4.38.2 requires tokenizers<0.19,>=0.14, but you have tokenizers 0.21.0 which is incompatible.[0m[31m
[0mSuccessfully installed numpy-1.26.4
```

```python
import torch
```

<!-- grade_id: cell-e0aba96eeaf8460e, points: 0, grade: false, solution: true -->
```python
step_size=0.1
def minimize(x0, num_iters):
    # Make x a tensor that requires gradients
    x = x0.clone().detach().requires_grad_(True)
    
    for _ in range(num_iters):
        f = f_torch(x)
        f.backward()
        x.data = x.data - step_size * x.grad
        x.grad.zero_()
    
    return x
```

<!-- grade_id: cell-3453f58f7aaffae2, points: 1, grade: true, solution: false -->
```python
x_opt = minimize(torch.zeros(3), 30)
assert (x_opt - torch.Tensor([0.1667, 0.0000, 0.0000])).norm() < 1e-3
```

Output:
```
Cell In[9], line 1
----> 1 x_opt = minimize(torch.zeros(3), 30)
      2 assert (x_opt - torch.Tensor([0.1667, 0.0000, 0.0000])).norm() < 1e-3

Cell In[8], line 7, in minimize(x0, num_iters)
      4 x = x0.clone().detach().requires_grad_(True)
      6 for _ in range(num_iters):
----> 7     f = f_torch(x)
      8     f.backward()
      9     x.data = x.data - step_size * x.grad

NameError: name 'f_torch' is not defined
```

### Problem 2: Working with PyTorch datasets
Find the mean brightness of CIFAR-10 images (https://www.cs.toronto.edu/~kriz/cifar.html) by iterating through the dataset and summing the pixel intensities.

```python
!pip install torchvision
```

Output:
```
Requirement already satisfied: torchvision in /opt/anaconda3/envs/datasci507/lib/python3.12/site-packages (0.17.2)
Requirement already satisfied: numpy in /opt/anaconda3/envs/datasci507/lib/python3.12/site-packages (from torchvision) (1.26.4)
Requirement already satisfied: torch==2.2.2 in /opt/anaconda3/envs/datasci507/lib/python3.12/site-packages (from torchvision) (2.2.2)
Requirement already satisfied: pillow!=8.3.*,>=5.3.0 in /opt/anaconda3/envs/datasci507/lib/python3.12/site-packages (from torchvision) (11.1.0)
Requirement already satisfied: filelock in /opt/anaconda3/envs/datasci507/lib/python3.12/site-packages (from torch==2.2.2->torchvision) (3.18.0)
Requirement already satisfied: typing-extensions>=4.8.0 in /opt/anaconda3/envs/datasci507/lib/python3.12/site-packages (from torch==2.2.2->torchvision) (4.12.2)
Requirement already satisfied: sympy in /opt/anaconda3/envs/datasci507/lib/python3.12/site-packages (from torch==2.2.2->torchvision) (1.13.3)
Requirement already satisfied: 
```

```python
import torch
import numpy as np
from torchvision import datasets
import torchvision.transforms as transforms


trainset = datasets.CIFAR10(root='./data',
                            train=True,
                            download=True,
                            transform=transforms.ToTensor())
```

Output:
```
File /opt/anaconda3/envs/datasci507/lib/python3.12/site-packages/torch/_dynamo/__init__.py:2
      1 import torch
----> 2 from . import allowed_functions, convert_frame, eval_frame, resume_execution
      3 from .backends.registry import list_backends, lookup_backend, register_backend
      4 from .code_context import code_context

File /opt/anaconda3/envs/datasci507/lib/python3.12/site-packages/torch/_dynamo/allowed_functions.py:30
     28 from . import config
     29 from .external_utils import is_compiling
---> 30 from .utils import hashable, is_safe_constant, NP_SUPPORTED_MODULES
     32 """
     33 A note on allowed functions:
     34 
   (...)
     47 skipfiles" there.
     48 """
     51 class FunctionIdSet:

ImportError: cannot import name 'NP_SUPPORTED_MODULES' from 'torch._dynamo.utils' (/opt/anaconda3/envs/datasci507/lib/python3.12/site-packages/torch/_dynamo/utils.py)
```

```python
from matplotlib import pyplot as plt


def imshow(img):
    npimg = img.numpy()
    plt.imshow(np.transpose(npimg, (1, 2, 0)))
    plt.show()


for i, (img, label) in enumerate(trainset):
    if i >= 5:  # Stop after 5 images
        break
    imshow(img)
    print(f"Label: {label}")
```

<!-- grade_id: cell-9f3da1d5dcc25501, points: 0, grade: false, solution: true -->
```python
def mean_brightness_directly(dataset):
    total_brightness = 0.0
    num_images = 0
    
    # Iterate through all images in the dataset
    for image, _ in dataset:
        # image shape is (3, 32, 32)
        # Take mean across all values (channels and pixels)
        total_brightness += image.mean().item()
        num_images += 1
    
    # Calculate mean brightness across all images
    mean_brightness = total_brightness / num_images
    
    return mean_brightness
```

```python
def mean_brightness_directly(dataset):
    # Create a DataLoader to load all images at once
    dataloader = torch.utils.data.DataLoader(dataset, batch_size=len(dataset), shuffle=False)
    
    # Get the single batch containing all images
    images, _ = next(iter(dataloader))
    
    # Calculate mean across all dimensions: batch, channels, height, width
    return images.mean().item()
```

<!-- grade_id: cell-321bdf2544b88758, points: 1, grade: true, solution: false -->
```python
assert abs(mean_brightness_directly(trainset) - 0.4733) < 1e-4
```

### Problem 3: Data loaders
Find the mean brightness of CIFAR-10 images by iterating through the dataset using a DataLoader with a batch size of 256.

<!-- grade_id: cell-36a333d71d7c1207, points: 0, grade: false, solution: true -->
```python
from torch.utils.data import DataLoader

def mean_brightness_dataloader(dataset):
    loader = DataLoader(dataset, batch_size=256, shuffle=True)
    
    total_brightness = 0.0
    total_images = 0
    
    # Iterate through batches
    for images, _ in loader:
        total_brightness += images.mean() * len(images)
        total_images += len(images)
    
    # Calculate final mean
    return (total_brightness / total_images).item()
```

<!-- grade_id: cell-82c31d17f7ec115a, points: 1, grade: true, solution: false -->
```python
assert abs(mean_brightness_dataloader(trainset) - 0.4733) < 1e-4
```

### Problem 4: Minimizing a loss function using SGD
Find the mean brightness of CIFAR-10 images in two ways: by iterating through the dataset and summing the pixel intensities and with stochastic gradient descent, by minimizing squared error.

<!-- grade_id: cell-8f04dd4d34016148, points: 0, grade: false, solution: true -->
```python
def mean_brightness_sgd(dataset):
    loader = DataLoader(dataset, batch_size=256, shuffle=True)
    cur_mean = torch.zeros(1)
    cur_mean.requires_grad_(True)
    
    optimizer = torch.optim.SGD([cur_mean], lr=0.1)
    
    for epoch in range(1):
        for images, _ in loader:
            # Zero gradients
            optimizer.zero_grad()
            
            # Calculate MSE loss between cur_mean and batch mean
            batch_mean = images.mean()
            loss = (cur_mean - batch_mean) ** 2
            
            # Backpropagate and update
            loss.backward()
            optimizer.step()
    
    return cur_mean.data.item()
```

<!-- grade_id: cell-c70dae2beab0f068, points: 1, grade: true, solution: false -->
```python
assert abs(mean_brightness_sgd(trainset) - 0.4733) < 1e-2
```

```python

```